<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 0  -->
  <question type="category">
    <category>
        <text>$course$/1081_405899001_基礎程式設計的預設/串列</text>

    </category>
  </question>

<!-- question: 32492  -->
  <question type="coderunner">
    <name>
      <text>刪除串列中某數值元素</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div class="O1">
<div class="O1">
<div class="O1">
<div class="O1">
<div class="O1">請撰寫一個函式，名為 removeElement，它有兩個參數nums(為串列型別)及val(整數型別)。</div>
<div class="O1">請直接從nums串列中將所有數值為val的元素刪除，並傳回刪除val元素後的nums串列還剩餘有多少元素(即長度為多少)。</div>
<div class="O1">傳回的數值與刪除元素後nums內容如範例所示。</div>
<div class="O1"></div>
<div class="O1">提示：可使用串列的remove( <em>value</em> )與 count( <em>value </em>)函式協助完成本問題。</div>
</div>
</div>
</div>
<div class="O1"></div>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>30.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>2</precheck>
    <showsource>0</showsource>
    <answerboxlines>8</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def removeElement(nums, val):
    for i in range(nums.count(val)):
        nums.remove(val)
    return len(nums)</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>0</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>nums = [3,2,2,3]
print(removeElement(nums,3),nums)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>2 [2, 2]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>nums = [0,1,2,2,3,0,4,2]
print(removeElement(nums,2),nums)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>5 [0, 1, 3, 0, 4]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>nums = [3,2,2,3]
print(removeElement(nums,1),nums)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>4 [3, 2, 2, 3]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>nums = [3,3,3]
print(removeElement(nums,3),nums)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>0 []</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>nums = []
print(removeElement(nums,0),nums)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>0 []</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 32490  -->
  <question type="coderunner">
    <name>
      <text>建立不重複元素串列的函式</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div class="O1">
<div class="O1">
<div class="O1">請撰寫一個函式名為 unique，有一個型別為串列的參數L。</div>
<div class="O1">此函式不會修改L的內容，但會傳回一個新串列，</div>
<div class="O1">其內容為所有L裡<strong>不重複</strong>的元素。</div>
</div>
<div class="O1"></div>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>30.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>2</precheck>
    <showsource>0</showsource>
    <answerboxlines>8</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def unique(L):
    newL = []
    for e in L:      
        if e not in newL:
            newL.append(e)
    return newL</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>0</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>unique_L=unique([1,2,2,2,3])
unique_L.sort()
print(unique_L)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[1, 2, 3]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>unique_L=unique([1,2,3,4,5])
unique_L.sort()
print(unique_L)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[1, 2, 3, 4, 5]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>unique_L=unique([])
unique_L.sort()
print(unique_L)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>unique_L=unique([True,True,True,True,True])
unique_L.sort()
print(unique_L)</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>[True]</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text><![CDATA[unique_L=unique(["4","3","2","3","1"])
unique_L.sort()
print(unique_L)]]></text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>['1', '2', '3', '4']</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
    </testcases>
  </question>

<!-- question: 32491  -->
  <question type="coderunner">
    <name>
      <text>比對兩串列是否有相同的不重複元素</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div class="O1">
<div class="O1">
<div class="O1">
<div class="O1">請撰寫一個函式名為 common，它有兩個參數L1及L2，型別皆為串列。</div>
<div class="O1">如果所有在L1裡不重複的元素，都存在於L2，</div>
<div class="O1">而且所有在L2裡不重複的元素，也都存在於L1，</div>
<div class="O1">則傳回True，否則傳回False。</div>
<div class="O1"></div>
<div class="O1">你可以使用前一題的unique函式（題目不提供需自行撰寫）協助完成本問題。</div>
</div>
</div>
<div class="O1"></div>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>40.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>2</precheck>
    <showsource>0</showsource>
    <answerboxlines>8</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload></answerpreload>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer>def unique(L):
    newL = []
    for e in L:      
        if e not in newL:
            newL.append(e)
    return newL

def common(L1,L2):
    unique_L1=unique(L1)
    unique_L2=unique(L2)
    if len(unique_L1) != len(unique_L2):
        return False
    else:
        for e in unique_L1:
            if e not in unique_L2:
                return False
        return True</answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>0</hoisttemplateparams>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000" >
      <testcode>
                <text>print(common([1,2,3],[3,1,2]))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>print(common([1,1,1],[1]))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>SHOW</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>print(common([1],[1,2]))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>False</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>print(common([1,3,3,5],[2,4,4,6]))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>False</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="1" mark="1.0000000" >
      <testcode>
                <text>print(common([],[]))</text>
      </testcode>
      <stdin>
                <text></text>
      </stdin>
      <expected>
                <text>True</text>
      </expected>
      <extra>
                <text></text>
      </extra>
      <display>
                <text>HIDE</text>
      </display>
    </testcase>
    </testcases>
  </question>

</quiz>